## Problem 1: De-duplicate a list

Given the following list,

    data=[5,4,6,1,9,0,3,9,2,7,10,8,4,7,1,2,7,6,5,2,8,2,0,1,1,1,2,10,6,2]

remove all the duplicated elements from this list. Note that the **order is preserved.** The order proceeds from smallest list index to greatest (i.e., from left-to-right).

<!-- grade_id: cell-c4e635caf51a3066, points: 0, grade: false, solution: true -->
```python
def uniq(input):
    '''
    this function remove all the duplicated elements from the input list
    inputs: input: an integer list
    '''
    output = []
    for i in input:
        if i not in output:
            output.append(i)
    return output
    

data=[5,4,6,1,9,0,3,9,2,7,10,8,4,7,1,2,7,6,5,2,8,2,0,1,1,1,2,10,6,2]

answer = uniq(data) 
```

<!-- grade_id: cell-963123c9c0289538, points: 5, grade: true, solution: false -->
```python
assert answer == [5,4,6,1,9,0,3,2,7,10,8]
```

## Problem 2: Compute the sum

Write a function that computes the sum of all non-negative integers (i.e. >=0) up to
and including a specified value, `n`.

**Hint**: Remember to use all the input/output validation checking techniques we discussed.

<!-- grade_id: cell-29d29d5b671acb58, points: 0, grade: false, solution: true -->
```python
def compute_sum_to_n(n):
    '''
    this function computes the sum of all non-negative integers up to and including a specified value n
    inputs: n: a non-negative integer number
    '''
    sum = 0
    i = 1
    while i <= n:
        sum += i
        i += 1
        
    return sum
```

<!-- grade_id: cell-76fa0090a0c926ab, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-8323d03ae4e1aad1, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-8ee2897af4252472, points: 2, grade: true, solution: false -->
```python
assert len(compute_sum_to_n.__doc__)
```

## Problem 3: All sublists

Generate all sublists of a list. Note there will be `2**len(x)` subsets of the
list x. For example, the list `[0,1,2]` has seven sublists:
`[[0],[1],[2],[0,1],[1,2],[0,2],[0,1,2]]`. Note that you don't have to mind the
ordering of each sublist.  Also, we don't care about the empty list (`[]`). The
function should accept input consisting of only a list of **unique** items.

* Hint: use recursion.
* Hint: You probably want to use `list` instead of `set`.

Here is the output of `all_sublists(range(4))`:

    [[0],
     [1],
     [2],
     [3],
     [2, 3],
     [1, 2],
     [1, 3],
     [1, 2, 3],
     [0, 1],
     [0, 2],
     [0, 3],
     [0, 2, 3],
     [0, 1, 2],
     [0, 1, 3],
     [0, 1, 2, 3]]

<!-- grade_id: cell-8d4d4b4a52f1bf1e, points: 0, grade: false, solution: true -->
```python
# from itertools import chain, combinations
def all_sublists(x):
    '''
    this function generates all sublists of a list except the empty sublist.
    inputs: x: a list that should only consists of unique items
    '''
    assert isinstance(x,list)
    assert len(x) == len(set(x))
    
    if x == []:
        return [x]

    sub = all_sublists(x[1:])

    return sub+[[x[0]]+i for i in sub]
    
#   return chain.from_iterable(combinations(x,n) for n in range(len(x)+1))
```

<!-- grade_id: cell-cfefcb73cddb53fc, points: 20, grade: true, solution: false -->
```python
### BEGIN  TESTS

ans = [[3], [2], [0], [2, 0], [3, 2], [3, 0], [3, 2, 0]]
assert 0 == len(set(map(frozenset,ans)) - set(map(frozenset,all_sublists([3,0,2]))))
ans = [[1], ['a'], [3], ['a', 3], [1, 'a'], [1, 3], [1, 'a', 3]]
assert 0 == len(set(map(frozenset,ans)) - set(map(frozenset,all_sublists([1,3,'a']))))

### END  TESTS
```

<!-- grade_id: cell-9d78fb31701077c6, points: 3, grade: true, solution: false -->
```python
import pytest
with pytest.raises(AssertionError):
    all_sublists(3)
```

<!-- grade_id: cell-9d78fb31701077c5, points: 3, grade: true, solution: false -->
```python
with pytest.raises(AssertionError):
    all_sublists([1,1])

```

<!-- grade_id: cell-c4fa5bbf21333553, points: 3, grade: true, solution: false -->
```python
assert len(all_sublists.__doc__)
```